# Tokenization

This notebook introduces some of the basic tools and idea for working with natural language (text), namely tokenization. Tokenization is the process of turning text into a sequence of words, with punctuation and common (stop) words removed.

## Imports

```python
import types
```

## Tokenization

```python
PUNCTUATION = '`~!@#$%^&*()_-+={[}]|\:;"<,>.?/}\t\n'
```

Write a generator function, `remove_punctuation`, that removes punctuation from an iterator of words and yields the cleaned words:

* Strip the punctuation characters at the beginning and end of each word.
* Replace the character `-` by a space if found in the middle of the word and split on that white space to yield multiple words.
* If a word is all punctuation, don't yield it at all.

<!-- grade_id: tokenizationa, points: 0, grade: false, solution: true -->
```python
def remove_punctuation(words, punctuation=PUNCTUATION):
    """Remove punctuation from an iterator of words, yielding the results."""
    print(words)
    arr = []
    stripF = 0;
    stripB = len(words);
    talk = words[0][stripF:stripB]
    print(talk)
    print(words[0])
```

<!-- grade_id: tokenizationb, points: 2, grade: true, solution: false -->
```python
assert list(remove_punctuation(['!data;']))==['data']
assert list(remove_punctuation(['!data-science:']))==['data', 'science']
assert list(remove_punctuation(['!!']))==[]
assert isinstance(remove_punctuation(['!!']), types.GeneratorType)
```

Output:
```
['!data;']
!
!data;
TypeError                                 Traceback (most recent call last)
<ipython-input-18-78edb0dc0caa> in <module>()
----> 1 assert list(remove_punctuation(['!data;']))==['data']
      2 assert list(remove_punctuation(['!data-science:']))==['data', 'science']
      3 assert list(remove_punctuation(['!!']))==[]
      4 assert isinstance(remove_punctuation(['!!']), types.GeneratorType)

TypeError: 'NoneType' object is not iterable
```

Write a generator function, `lower_words`, that makes each word in an iterator lowercase, yielding each lowercase word:

<!-- grade_id: tokenizationc, points: 0, grade: false, solution: true -->
```python
def lower_words(words):
    """Make each word in an iterator lowercase."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: tokenizationd, points: 2, grade: true, solution: false -->
```python
assert isinstance(lower_words('AAA'), types.GeneratorType)
assert list(lower_words('This IS NOT LoWerCaSe'.split(' ')))==['this', 'is', 'not', 'lowercase']
```

[Stop words](https://en.wikipedia.org/wiki/Stop_words) are common words in text that are typically filtered out when performing natural language processing. Typical stop words are *and*, *of*, *a*, *the*, etc.

Write a generator function, `remove_stop_words`, that removes stop words from an iterator, yielding the results:

<!-- grade_id: tokenizatione, points: 0, grade: false, solution: true -->
```python
def remove_stop_words(words, stop_words=None):
    """Remove the stop words from an iterator of words.
    
    stop_words can be provided as a list of words or a whitespace separated string of words.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: tokenizationf, points: 2, grade: true, solution: false -->
```python
assert list(remove_stop_words('the begin to the end a of the day'.split(' '), stop_words='a the')) == \
    ['begin', 'to', 'end', 'of', 'day']
assert list(remove_stop_words('the begin to the end a of the day'.split(' '), stop_words=['a', 'the'])) == \
    ['begin', 'to', 'end', 'of', 'day']
assert list(remove_stop_words('the begin to the end a of the day'.split(' '))) == \
    ['the', 'begin', 'to', 'the', 'end', 'a', 'of', 'the', 'day']
```

[Tokenization](https://en.wikipedia.org/wiki/Lexical_analysis#Tokenization) is the process of taking a string or line of text and returning a sequence of words, or *tokens*, with the following transforms applied

* Punctuation removed
* All words lowercased
* Stop words removed

Write a generator function, `tokenize_line`, that yields tokenized words from an input line of text. 

<!-- grade_id: tokenizationg, points: 0, grade: false, solution: true -->
```python
def tokenize_line(line, stop_words=None, punctuation=PUNCTUATION):
    """Split a string into a list of words, removing punctuation and stop words."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: tokenizationh, points: 2, grade: true, solution: false -->
```python
assert isinstance(tokenize_line("This, is the way; that things will end"), types.GeneratorType)
assert list(tokenize_line("This, is the way; that things will end", stop_words=['the', 'is'])) == \
    ['this', 'way', 'that', 'things', 'will', 'end']
```

Write a generator function, `tokenize_lines`, that can yield the tokens in an iterator of lines of text.

<!-- grade_id: tokenizationi, points: 0, grade: false, solution: true -->
```python
def tokenize_lines(lines, stop_words=None, punctuation=PUNCTUATION):
    """Tokenize an iterator of lines, yielding the tokens."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: tokenizationj, points: 2, grade: true, solution: false -->
```python
wasteland = """
APRIL is the cruellest month, breeding
Lilacs out of the dead land, mixing
Memory and desire, stirring
Dull roots with spring rain.
"""

assert isinstance(tokenize_lines(wasteland.splitlines()), types.GeneratorType)

assert list(tokenize_lines(wasteland.splitlines(), stop_words='is the of and')) == \
    ['april','cruellest','month','breeding','lilacs','out','dead','land',
     'mixing','memory','desire','stirring','dull','roots','with','spring',
     'rain']
```

## Tokenize song lyrics

Now use all of the above functions to perform tokenization on the set of song lyrics from this Kaggle hosted dataset:

https://www.kaggle.com/mousehead/songlyrics

* You should be able to perform this in a memory efficient manner.
* Read your stop words from the included `data/stopwords.txt` file.

Here is the dataset loaded as a Pandas `DataFrame`:

```python
import pandas as pd
df = pd.read_csv("/data/songlyrics/songdata.csv")
df.head()
lyrics = df['text']
```

If we extract the `text` column, we get an iterator over the lyrics. **Remember each lyric can and will have multiple lines.** Here is the total number of lyrics:

```python
len(lyrics)
```

Read the file `data/stopwords.txt` and tokenize the file into a list of stop words:

<!-- grade_id: tokenizationm, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: tokenizationn, points: 2, grade: true, solution: false -->
```python
assert len(stop_words)==174
assert type(stop_words)==list
```

```python
print(stop_words)
```

Now iterate through the lyrics and for each lyric:

1. Split the lyric into lines.
2. Call `tokenize_lines` on the lyric eliminating the above stop words (and punctuation).
3. Count the total number of words across all lyrics (excluding stop words) and save the result in a variable named `nwords`.

If you do all of this in a memory efficient manner, the total memory consumption of this notebook shouldn't go over around 250MB. Most of that is just using Pandas to read the dataset into memory. If you construct a full list of all the words and *then* count them all, your memory consumption will be 3-4x that. This should only take a few minutes to run.

<!-- grade_id: tokenizationp, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print("Total number of words: {}".format(nwords))
```